```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

# Q1 SSH, command line, exit codes, etc 

SSH to your account at `10.4.66.86` (e.g. `ssh your-user-name@10.4.66.86`) and follow the puzzle starting in the directory `/schnitzeljagd`. Place the final answer in the variable `final_answer`.

<!-- grade_id: cell-d6b141313f4f3a09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3d5b7fe600c4c6ba, points: 1, grade: true, solution: false -->
```python
assert ads_hash(final_answer) == 'dfb81c4fac'
```

# Q2 Starting anaconda environment from command line

On the above server, we have anaconda installed at `/anaconda` so from the command line, you can start the anaconda `base` environment like this:

    source /anaconda/bin/activate
    
If you do this, your command-line prompt will start with `(base)` to indicate the name of your anaconda environment.

Once you have started an anaconda environment, you can run `conda env list` to see what environments are available. You can start an environment with:

    conda activate <environment-name>
    
Steps:

0) Login to `strawlab-course-server` as described above.
1) Start the environment for this course.
2) Change to the directory for this course and exercise in `/courses`.
3) Run the python program `program1.py`
4) Put the program console output into the variable `program1_output` below

<!-- grade_id: cell-6d9ed8460134bae3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc67d9f37b59741d, points: 0, grade: true, solution: false -->
```python
assert ads_hash(program1_output) == '983bd614bb'
```