# Getting Started with Jupyter

```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

```python
1+1
```

Output:
```
2
```

# Q1 first question

Assign to the variable `x` the result of `1 + 1`.

<!-- grade_id: cell-ffc933aca86d4cfd, points: 0, grade: false, solution: true -->
```python
x = 1 + 1
```

```python
assert x == 2
```

<!-- grade_id: cell-086097d6cabebf8c, points: 1, grade: true, solution: false -->
```python
assert ads_hash(x) == 'd4735e3a26'
```

```python
x = 10
```